# Computer Assignment 1

***
**Task 1**:  Simple plotting of signals
***

Write a general function that will generate values from a finite-length sinusoid of the form
$$
  x[n] = A \cos(\omega_0 n + \phi) 
$$
The function will need four input arguments: three for the parameters (scalar floats $A$, $\omega_0$, and $\phi$), and one vector containing the values of $n$ for which $x[n]$ must be calculated. The function should take the form

> `xv = getcosinevals(A, w0, phi, nv)`

With the obvious correspondence between parameters and variables. 

Test the function by generating and plotting values of the signal
$$
  x_1[n] = 2 \cos \left( \frac{2\pi}{12} n + \frac{\pi}{4} \right)
$$
over the range $-10$ to $15$.  The horizontal axis should extend only over the range indicated, and the axes should be labelled accordingly.  Each signal should be displayed using a stem plot since it's a discrete signal.

```python
import numpy as np
import matplotlib.pyplot as plt

def getcosinevals(A, w0, phi, nv):
    ...
    return xv

...
n1v = ...
x1v = getcosinevals(A, w0, phi, n1v)

# Stem plot with labels
plt.stem(n1v, x1v)
plt.xlabel('$n$')
plt.ylabel('$x_1[n]$')

# Make axes a bit prettier
ax = plt.gca()
ax.autoscale(enable=True, axis='x', tight=True)
```

***
**Task 2**:  Simple filtering, with input and output plotted on same axes
***

A lowpass filter has an impulse response    
$$\begin{aligned}
  h[n] = \begin{cases}
    \frac{1}{7} \qquad & 0 \leq n \leq 6 \\
    0 \qquad & \text{otherwise}.
  \end{cases}
\end{aligned}$$
Suppose the input to the filter is the (one-sided) sinusoidal sequence $$x[n] = u[n] \cos(\omega_0 n),$$ with $\omega_0 = \pi/10$, and the output is $y[n]$. Calculate and plot the input and output sequence on the same set of axes over the range $-10, \ldots, 70$.  Note that after the initial transient dies down, the output signal is also sinusoidal. What is the magnitude and phase of this output relative to the input?

```python
import numpy as np
import matplotlib.pyplot as plt

# The block of code below should populate xv with the required x[n] values for each corresponding element of nv
n2v = np.arange(-10,71)  # vector of n values to consider
...
x2v = ...
plt.stem(n2v,x2v,linefmt='r');

# The block of code below should populate yv with the required y[n] values for each corresponding element of nv.
# Initially you should calculate these values using two nested for loops, but subsequently you can explore using
# numpy.convolve
...
y2v = ...
plt.stem(n2v,y2v,linefmt='b');

plt.xlabel("n")
plt.ylabel("Signal value")
plt.legend(["x[n]","y[n]"]);
```

***
**Task 3**:  Frequency response
***

Write a function that takes in a right-sided filter impulse response $h[n]$ and a set of frequencies $\omega$, and returns the corresponding values of $H(e^{j \omega})$.  The function declaration may look as follows:

> `Hv = freqresp(hv, wv)`

Here `hv` is a vector with $p$ elements specifying all the nonzero values of the impulse response, such that
`hv[0]`=$h[0]$, `hv[1]`=$h[1]$, and so on. The input `wv` is a vector of real-valued frequency values, and the output `Hv` must contain the complex frequency response values corresponding to the elements of `wv`.

Use this function to find the frequency response values of the moving average lowpass filter
$$\begin{aligned}
  h_3[n] = \begin{cases}
    \frac{1}{16} (n + 1) \qquad & 0 \leq n \leq 3 \\
    \frac{1}{16} (7 - n) \qquad & 4 \leq n \leq 6 \\
    0 \qquad & \text{otherwise}.
  \end{cases}
\end{aligned}$$
for a range of frequencies $\omega = 0, 0.01, \ldots, 4\pi$. Plot the magnitude `abs(H)` as a function of $\omega$. On a separate set of axes, plot the phase `angle(H)` as a function of $\omega$. Interpret the results you obtain, and note the periodicity of the frequency response. What effect will the system have on an input complex exponential with frequency $\omega = \pi/2$?

```python
import numpy as np
import matplotlib.pyplot as plt

def freqresp(hv, wv):
    Hv = np.zeros(wv.shape, dtype=np.cdouble)
    ...
    return Hv
    
# The block below should generate the vector of values htv corresponding to the elements of nv
n3v = np.arange(0,7)
h3v = ...
...

# Plot impulse response
plt.figure
plt.stem(n3v,h3v)
plt.xlabel('$n$')
plt.ylabel('$h_3[n]$')
ax = plt.gca()
ax.set_title('Filter impulse response')

# Vector of values of omega at which to calculate frequency response
w3v = np.arange(0, 4*np.pi, 0.01)

# Calculate frequency response values
H3v = freqresp(h3v, w3v)

# Plot frequency response
fig, axs = plt.subplots(2, 1, constrained_layout=True)

axs[0].plot(w3v, np.abs(H3v))
axs[0].set_title('Filter frequency response')
axs[0].set_ylabel('$|H(e^{j \omega})|$')

axs[1].plot(w3v, np.angle(H3v))
axs[1].set_xlabel('$\omega$')
axs[1].set_ylabel('$\t{angle}~H(e^{j \omega})$');
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. Export as a PDF and submit. **Please save before exporting!**

Tests for lab_ass01_q1a (points: None)
```
>>> assert len(n1v) == 15 - -10 + 1
```
failure_message: len(n1v) does not return correct length
success_message: len(n1v) returns correct length
```
>>> assert n1v.shape == x1v.shape
```
failure_message: x1v does not have same shape as n1v
success_message: x1v has same shape as n1v


Tests for lab_ass01_q1b (points: None)
```
>>> assert n2v.shape == x2v.shape
```
failure_message: xv does not have same shape as nv
success_message: xv has same shape as nv
```
>>> assert all(x2v[n2v < 0] == 0)
```
failure_message: elements of xv are not all zero where nv<0
success_message: elements of xv are all zero where nv<0
```
>>> assert n2v.shape == y2v.shape
```
failure_message: y2v does not have same shape as n2v
success_message: y2v has same shape as n2v
```
>>> assert all(x2v[n2v < 0] == 0)
```
failure_message: elements of yv are not all zero where nv<0
success_message: elements of yv are all zero where nv<0


Tests for lab_ass01_q1c (points: None)
```
>>> assert n3v.shape == h3v.shape
```
failure_message: h3v does not have same shape as n3v
success_message: h3v has same shape as n3v
```
>>> assert w3v.shape == H3v.shape
```
failure_message: H3v does not have same shape as w3v
success_message: H3v has same shape as w3v
